# Homework 2: DataFrames, Data Visualization, and Functions

## Due Monday, April 20th at 11:59PM

Welcome to Homework 2! This week, we will cover DataFrame manipulations, making visualizations, and defining functions. You can find additional help on these topics in  [BPD 6, 9-12](https://notes.dsc10.com/01-getting_started/functions-defining.html) in the `babypandas` notes and [CIT 7-7.3](https://inferentialthinking.com/chapters/07/Visualization.html) in the textbook.

### Instructions

Remember to start early and submit often. You are given six slip days throughout the quarter to extend deadlines. See the syllabus for more details. With the exception of using slip days, late work will not be accepted unless you have made special arrangements with your instructor.

**Important**: For homeworks, the `otter` tests don't usually tell you that your answer is correct. More often, they help catch careless mistakes. It's up to you to ensure that your answer is correct. If you're not sure, ask someone (not for the answer, but for some guidance about your approach). These are great questions for office hours (see the schedule on the [Calendar](https://dsc10.com/calendar)) or Campuswire. Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. 

**Please do not use for-loops for any questions in this homework.** If you don't know what a for-loop is, don't worry – we haven't covered them yet. But if you do know what they are and are wondering why it's not OK to use them, it is because loops in Python are slow, and looping over arrays and DataFrames should usually be avoided.

<font color=red>**🚨 If you create a data visualization that is too cluttered to read or takes more than a few seconds to generate, this is a sign you are doing something wrong. Do not submit code like this, or the Gradescope autograder may fail to grade your entire assignment. It's better to leave a question blank than to submit code that will cause the autograder to fail on your full assignment.**</font>

```python
# Please don't change this cell, but do make sure to run it
import babypandas as bpd
import numpy as np
np.set_printoptions(legacy='1.21')

import matplotlib.pyplot as plt
plt.style.use('seaborn-v0_8-colorblind')
plt.rcParams['figure.figsize'] = (10, 5)

import otter
grader = otter.Notebook()
```

## 1. See ya at IKEA! 🪑 🧆 

IKEA is a popular furniture store known for its affordable furniture, maze-like store layout, and their food court's Swedish meatballs. The first IKEA store opened in Sweden in 1958, and they've since expanded to have over 470 IKEA stores globally, including one in San Diego.

The file named `ikea.csv` in the `data/` directory contains data on various IKEA products. This data includes information on each product's price, description, and size, among other attributes. Products with unknown `'depth'`, `'height'`, and `'width'` values have these attributes recorded as `0.0`.

```python
ikea_full = bpd.read_csv('data/ikea.csv')
ikea_full
```

**Question 1.1.** We will only be using the `'name'`, `'item_id'`, `'category'`, `'price'`, `'short_description'`, and `'height'` columns from the `ikea_full` DataFrame. Create a new DataFrame called `ikea` that contains all the information from the `ikea_full` DataFrame, but with only the 6 columns mentioned here.

```python
ikea = ...
ikea
```

```python
grader.check("q1_1")
```

### Aside: String methods and Series

**Before proceeding to Question 1.2, make sure to read this note.**

As we saw in [Lecture 3](https://dsc10.com/resources/lectures/lec03/lec03.html#String-methods), strings have several methods available to them. 

```python
'NORDVIKEN'.lower()
```

```python
'NORDVIKEN'.replace('O', '🎯')
```

`babypandas` makes it easy for us to use string methods on not just a single string, but a *Series* of strings. To do this, we follow the pattern

```py
ser.str.method(args)
```

where `ser` is a Series, `method` is the name of a method, and `args` are the arguments to that method (if any). The key difference between the usage of string-specific Series methods and other Series methods (like `ser.mean()`) is that we need to write `.str` first.

For example, the following expression returns a Series containing the values in the `name` column but converted to lowercase.

```python
ikea.get('name').str.lower()
```

The following expression replaces all `'O'`s with `'🥝'` in the index. (Yes, you can use `.str` with the index as well!)

```python
example = ikea.set_index('name')
example.index.str.replace('O', '🥝')
```

The following expression returns a Boolean Series, with `True` for products whose `short_description` contains the string `'desk'` and `False` for all other products. (`contains` is not a regular string method in Python, while the other examples – `lower` and `replace` – are.)

```python
ikea.get('short_description').str.contains('desk')
```

Remember that we use Boolean Series to make queries! If we query with this Boolean Series, we'll obtain a DataFrame of only the furniture items that have `'desk'` as part of the short_description.

```python
ikea[ikea.get('short_description').str.contains('desk')]
```

Now you're ready to proceed to Question 1.2. If you'd like another explanation of Series string methods, and `.str.contains` in particular, run the cell below to watch a video from a past iteration of DSC 10.

```python
from IPython.display import YouTubeVideo
id = 'TCcEhVA6Euw?si'
YouTubeVideo(id = id, height = "300", width = "600")
```

**Question 1.2.** Arya is interested in purchasing a shelf from IKEA. The `'short_description'` column in the `ikea` DataFrame contains an overview of the product, such as if the product is table or a shelving unit. To make it easier for Arya to see all of the possible product options that he can buy, assign `shelves` to a DataFrame that includes all the products that have the word `'shelving'` in their description, with any capitalization, sorted in ascending order of `'price'`.

***Hint:*** The product descriptions are strings, so they may have inconsistencies in how they're capitalized. If we want to include `'shelving'`, `'Shelving'`, `'ShElViNg'`, and any other variations in capitalization, what operation should we call on the description column **first**? (You may end up using `.str` twice!)

```python
# For a multi-step problem, it's helpful to define intermediate variables. 
# Feel free to do that here, or for any problem!

shelves = ...
shelves
```

```python
grader.check("q1_2")
```

To see what types of shelving units there are, run the following cell, which uses the Series method `unique()`.

```python
shelves.get('category').unique()
```

**Question 1.3.** After looking through the possible shelving options, Arya decides that he might as well purchase a shelving product that can also function as a table or desk! However, he heard from Oren that the `'EKET'` table and desk isn't very good, so he wants to avoid that one. Assign `no_eket` to a DataFrame containing all the products in `shelves` that are categorized as `'Tables & desks'` except for the one named `'EKET'`. `no_eket` should have the same columns as `shelves`.

```python
no_eket = ...
no_eket
```

```python
grader.check("q1_3")
```

**Question 1.4.** Arya doesn't want to buy the most expensive product but he also has enough money in his budget to afford buying something other than the cheapest product, so he won't buy the cheapest product either. Using the `no_eket` DataFrame, assign `final_choice` to the `'item_id'` of the product that he will buy under these conditions.  

```python
final_choice = ...
final_choice
```

```python
grader.check("q1_4")
```

**Question 1.5.** Next, Arya needs a place to sit. Plot a histogram that displays the distribution of the `height` of all products categorized as `'Chairs'` or `'Sofas & armchairs'`. Only include products for which we have height information available.

When plotting your histogram, remember to set `density=True` and `ec='w'`. You don't have to set the `bins` argument.

```python
# Create your histogram here.
...
```

**Question 1.6.** Among all products included in the histogram above, find the proportion of such products that are between 75cm and 100cm tall in `'height'`, inclusive. Assign this proportion to `prop`. Use the histogram to verify whether your answer looks correct, which may involve calculating some areas!

```python
prop = ...
prop
```

```python
grader.check("q1_6")
```

**Question 1.7.** After hours of wandering through IKEA, Arya has seen a lot of different kinds of furniture! He would like to know the median price of the products in each `'category'` to get a sense of how much a typical furniture item from each `'category'` should cost. Create a DataFrame called `typical_prices` where the index contains the `'category'` values in `ikea` and the only column is `'price'`, which contains the **median** price of all products in a given `'category'`.

```python
typical_prices = ...
typical_prices
```

```python
grader.check("q1_7")
```

**Question 1.8.** Set `ranked_categories` to an array of all the different `'category'` values in `ikea`, in ascending order of median price. 

***Hint:*** Use `np.array` to make sure `ranked_categories` is an array.



```python
ranked_categories = ...
ranked_categories
```

```python
grader.check("q1_8")
```

## 2. Broadway Bound 🎭

<center><img src='./images/hamilton.jpg' width = 400/></center>

In this problem, we'll analyze data on your favorite Broadway shows! We've provided a file, `./data/broadway.csv` that contains information about Broadway plays and musicals from 1985 to 2020. The dataset, which comes from [Kaggle](https://www.kaggle.com/datasets/jessemostipak/broadway-weekly-grosses?select=grosses.csv) and was originally scraped from [Playbill](https://playbill.com/grosses), contains a row for each week for each show on Broadway. For each week and show combination, we have the following pieces of information that we will use to generate insights about the state of musicals over time. 

|Column|Description|
|------|-----------|
|`'Week Ending'`|The week, referenced by the Sunday at the end of the week.|
|`'Show'`|The name of the play or musical.|
|`'Theatre'`|The location or venue at which the show was held.|
|`'Weekly Gross'`|The total revenue generated from ticket sales for that show during that week.|
|`'Seats Sold'`|The total number of seats sold for that show during that week.|
|`'Performances'`|The number of performances of the show held during that week, including "previews" before the official opening.|

```python
# Run this cell to load the dataset.
broadway = bpd.read_csv('data/broadway.csv') 
broadway
```

**Question 2.1.** Find the name of the show that generated the most revenue during the timespan of our dataset. Assign your answer to `most_revenue_show`. 

```python
most_revenue_show = ...
most_revenue_show
```

```python
grader.check("q2_1")
```

**Question 2.2.** Broadway shows come and go, but some stay longer than others. Which show in `broadway` was on Broadway for the greatest number of weeks (not necessarily consecutive)? Assign you answer to `most_weeks_show`.

```python
most_weeks_show = ...
most_weeks_show
```

```python
grader.check("q2_2")
```

**Question 2.3.** Add a column to `broadway` that contains, for each week and show combination, the average price of a ticket for that show during that week. Call your new column `'AvgTicket'`, and save the resulting DataFrame in the variable `broadway`.

```python
broadway = ...
broadway                    
```

```python
grader.check("q2_3")
```

**Question 2.4.** The musical *Wicked* has been at the Gershwin Theatre continuously since its opening in 2003. Create a line plot that displays how the average weekly ticket price for *Wicked* has changed over time.

```python
# Create your line plot here.
...
```

It looks like *Wicked* ticket prices fluctuate a lot throughout each year, which is referred to as *seasonal variation*. Most likely, this phenomenon is not particular to *Wicked* but exists to some degree for all Broadway shows. Let's try to find out what time of year Broadway tickets are most expensive. 

**Question 2.5.** Create a new DataFrame called `with_month` that has all the data in `broadway`, plus a column called `'Month'`. `'Month'` should contain strings with two characters (e.g. `'04'` for April) corresponding to the dates in the `'Week Ending'` column. 

```python
with_month = ...
with_month
```

```python
grader.check("q2_5")
```

**Question 2.6.** Create a line plot showing the mean value of `'AvgTicket'` for each month of the year, to see the seasonal variation. Then, assign `most_expensive` to the month of the year (again, as a string like `'04'`) where ticket prices tend to be the most expensive. Extract this month using Python code, but verify visually that it matches up with what your line plot shows.

```python
# Create your line plot here.
...
```

```python
most_expensive = ...
most_expensive 
```

```python
grader.check("q2_6")
```

**Question 2.7.** Which Broadway theatres have hosted the most performances? To answer this question, create a DataFrame named `theatre_proportions` indexed by `Theatre`, with one column `'Proportion'`. The `'Proportion'` column should contain, for each theatre, the proportion of performances in the dataset that were held in that theatre. Remember that each row of our dataset corresponds to multiple performances! Order the rows of `theatre_proportions` in descending order of `'Proportion'`. 

```python
theatre_proportions = ...
theatre_proportions
```

```python
grader.check("q2_7")
```

**Question 2.8.** Broadway shows tend to bring in a lot of money each week! With multiple performances, packed theatres, and high ticket prices, the revenue stream is flowing! Let's explore how much money was brought in each week by the show *Hamilton*,  a popular award-winning musical that debuted on Broadway in 2015. Plot a density histogram that shows the distribution of `'Weekly Gross'` for *Hamilton* over the period between its debut and the end of data collection. Use bins with a width of a quarter of a million dollars, and make sure that all the data from *Hamilton* is included within your bins.

```python
# Create your histogram here.
...
```

**Question 2.9.** Without writing any code, can you estimate the proportion of weeks for which *Hamilton* earned between \\$3,000,000 (inclusive) and \\$3,250,000 (exclusive)? Assign your answer to `prop_between` as a float between 0 and 1.

```python
prop_between = ...
prop_between
```

```python
grader.check("q2_9")
```

**Question 2.10.** Broadway shows often have interesting and memorable names that reflect their themes, characters, or settings. From classics like *The Phantom of the Opera* to modern hits like *Hamilton*, these names capture the imagination and hint at the unique stories waiting to be told on stage. How often do Broadway shows get one-word names? Assign `one_word` to the proportion of Broadway show names that are only one word. Remember, each show appears multiple times in our DataFrame!

***Note:*** Some shows are made up of numbers only, which are not really "words" in the conventional sense. We'll consider words to be chunks of text separated by spaces, so for example, the show `"1776"` should count as having one word.

```python
one_word = ...
one_word
```

```python
grader.check("q2_10")
```

## 3. Final Stretch 🧘‍♀️

Suppose we have a DataFrame named `data` with two numerical columns, `'x'` and `'y'`. Consider the following scatter plot, which was generated by calling `data.plot(kind='scatter', x='x', y='y')`:

<img src='images/q3_scatter.png' width=400/>

Now consider these two histograms:

<center>
    <table><tr>
        <td><center><b>Histogram A</b><br> <img src='images/q3_hist_one.png' width=400></center> </td>
        <td><center><b>Histogram B</b><br> <img src='images/q3_hist_two.png' width=400></center> </td>
    </tr></table>
</center>

**Question 3.1.** Which of the following lines of code generated Histogram A? Assign either `1`, `2`, `3`, or `4` to `which_code`.

 1. `data.plot(kind='hist', density=False, y='x')`
 2. `data.plot(kind='hist', density=False, y='y')` 
 3. `data.plot(kind='hist', density=True, y='x')`
 4. `data.plot(kind='hist', density=True, y='y')`

```python
which_code = ...
```

```python
grader.check("q3_1")
```

**Question 3.2.** Suppose we run this block of code:

```py
new_data = bpd.DataFrame().assign(
    x = data.get('x') / 5,
    y = data.get('y')
)
```
    
We then run 

`new_data.plot(kind='hist', density=True, y='x')`.

How will this histogram look compared to the histogram created by 

`data.plot(kind='hist', density=True, y='x')`, 

assuming both histograms are drawn on the same axes? Assign `histogram_difference` to either 1, 2, 3, or 4, corresponding to your choice.

1. The `new_data` histogram will be wider and taller than the `data` histogram.
2. The `new_data` histogram will be wider and shorter than the `data` histogram.
3. The `new_data` histogram will be narrower and taller than the `data` histogram.
4. The `new_data` histogram will be narrower and shorter than the `data` histogram.

_*Hint*_: Look at the end of [Lecture 7](https://dsc10.com/resources/lectures/lec07/lec07.html#Plotting-overlaid-histograms) for an example of two histograms drawn on the same axes.

```python
histogram_difference = ...
```

```python
grader.check("q3_2")
```

**Question 3.3.** Below, we show Histogram B again.

<img src="./images/q3_hist_two.png" width=400/>

What **percent** of values in Histogram B are between -4 (inclusive) and -2 (exclusive)? While we cannot answer this question exactly since we do not know where the bins start and end, we can still approximate the answer. Assign the variable `percent_between` to a number 1 through 5, corresponding to the closest answer.

1. 10% 
2. 13% 
3. 27%
4. 35%
5. 48%

```python
percent_between = ...
```

```python
grader.check("q3_3")
```

## Finish Line: Almost there, but make sure to follow the steps below to submit! 🏁

<hr style="color:Maroon;background-color:Maroon;border:0 none; height: 3px;">

**Citations:** Did you use any generative artificial intelligence tools to assist you on this assignment? If so, please tell us which tools you used by setting `gen_ai` to a **list** of numbers. Include all tools used, from the options below:

1. AI Tutor, in Tutor Mode
2. AI Tutor, in ChatGPT Mode
3. Any other outside generative AI tool, such as Gemini, Claude, ChatGPT, etc.

If you didn't use any such tools, set `gen_ai` to an empty list, `[]`.

```python
gen_ai = ...
```

```python
grader.check("citations")
```

<hr style="color:Maroon;background-color:Maroon;border:0 none; height: 3px;">

To submit your assignment:

1. Make sure to comment out the code in Question 2.1 that causes an error.
1. Select `Kernel -> Restart Kernel and Run All Cells` to ensure that you have executed all cells, including the test cells.
2. Read through the notebook to make sure everything is fine and all tests passed.
3. Run the cell below to run all tests, and make sure that they all pass.
4. Download your notebook as a `.ipynb` file using `File -> Download`, then upload your notebook  to Gradescope under "Lab 0".
5. **Stick around while the Gradescope autograder grades your work. Make sure you see that all tests have passed on Gradescope.**
6. Check that you have a confirmation email from Gradescope and save it as proof of your submission.

With homeworks, unlike with labs, the grade you see on Gradescope is **not your final score**. We will run correctness tests after the assignment's due date has passed.

```python
grader.check_all()
```

Tests for citations (points: None)
```
>>> isinstance(gen_ai, list)
True
```
failure_message: In Python, a list has square brackets and comma-separated entries, such as [1, 2, 3]


Tests for q1_1 (points: None)
```
>>> isinstance(ikea, bpd.DataFrame) and ikea.shape[1] == 6
True
```


Tests for q1_2 (points: None)
```
>>> isinstance(shelves, bpd.DataFrame) and shelves.shape[0] % 10 == 2
True
```
```
>>> shelves.index[4] - shelves.index[3] == 89
True
```
failure_message: Try again!


Tests for q1_3 (points: None)
```
>>> isinstance(no_eket, bpd.DataFrame) and no_eket.shape[1] == shelves.shape[1]
True
```
```
>>> len(no_eket.get('name').unique()) == 2
True
```
```
>>> no_eket.shape[0] == 3
True
```


Tests for q1_4 (points: None)
```
>>> final_choice in no_eket.get('item_id').values
True
```


Tests for q1_6 (points: None)
```
>>> int(prop * 10 ** 10 % 10 ** 3) == 931 or int(prop * 10 ** 10 % 10 ** 3) == 17
True
```
failure_message: Try again!
```
>>> np.isclose(prop, 0.3179317931793179) or np.isclose(prop, 0.5043630017452007)
True
```
failure_message: The correct answer is 0.504, but we originally had a mistake in the test, so we will accept both answers as correct.


Tests for q1_7 (points: None)
```
>>> isinstance(typical_prices, bpd.DataFrame) and typical_prices.shape == (17, 1)
True
```
```
>>> int(typical_prices.get('price').loc['Sofas & armchairs'] % 100) == 68
True
```
failure_message: Try again!


Tests for q1_8 (points: None)
```
>>> isinstance(ranked_categories, np.ndarray) and len(ranked_categories) == 17
True
```


Tests for q2_1 (points: None)
```
>>> bool(most_revenue_show in broadway.get('Show').unique())
True
```
failure_message: most_revenue_show should be the name of a show


Tests for q2_10 (points: None)
```
>>> bool(isinstance(one_word, float))
True
```


Tests for q2_2 (points: None)
```
>>> bool(isinstance(most_weeks_show, str))
True
```
```
>>> bool(most_weeks_show in broadway.get('Show').unique())
True
```
failure_message: most_weeks_show should be the name of a show


Tests for q2_3 (points: None)
```
>>> bool(broadway.shape[1] == 7)
True
```
```
>>> bool('AvgTicket' in broadway.columns)
True
```
failure_message: Check capitalization and spelling of your column name.


Tests for q2_5 (points: None)
```
>>> bool(broadway.shape[1] == 7 and with_month.shape[1] == 8 and ('Month' in with_month.columns))
True
```


Tests for q2_6 (points: None)
```
>>> bool(most_expensive in ['01', '02', '03', '04', '05', '06', '07', '08', '09', '10', '11', '12'])
True
```


Tests for q2_7 (points: None)
```
>>> bool(isinstance(theatre_proportions, bpd.DataFrame) and set(theatre_proportions.columns) == {'Proportion'})
True
```
```
>>> bool(np.isclose(theatre_proportions.get('Proportion').sum(), 1))
True
```
```
>>> bool(theatre_proportions.index[1] == 'Minskoff Theatre')
True
```
failure_message: Tests that the DataFrame is sorted.


Tests for q2_9 (points: None)
```
>>> bool(isinstance(prop_between, float))
True
```
```
>>> bool(prop_between >= 0 and prop_between <= 1)
True
```


Tests for q3_1 (points: None)
```
>>> bool(which_code in {1, 2, 3, 4})
True
```


Tests for q3_2 (points: None)
```
>>> bool(histogram_difference in {1, 2, 3, 4})
True
```


Tests for q3_3 (points: None)
```
>>> bool(percent_between in {1, 2, 3, 4, 5})
True
```
