# Lab 5:  Variability and the Normal Distribution

## Due Thursday, May 14th at 11:59PM

Welcome to Lab 5! In this lab you will practice calculating variance and standard deviation, and converting values to standard units. You will use these skills to compare grades in a course and you will use Chebyshev’s inequality to predict how hard the students should work in order to rank in the top 5% of the class. Finally, you will use confidence intervals to help college administrators plan for next quarter by predicting the enrollment in a new course. The topics from this lab are covered in the following readings:
* [CIT 14.2](https://www.inferentialthinking.com/chapters/14/2/Variability.html): Variability
* [CIT 14.3](https://www.inferentialthinking.com/chapters/14/3/SD_and_the_Normal_Curve.html): The Standard Deviation (SD) and the Normal Curve 
* [CIT 14.4](https://www.inferentialthinking.com/chapters/14/4/Central_Limit_Theorem.html): The Central Limit Theorem
* [CIT 14.5](https://www.inferentialthinking.com/chapters/14/5/Variability_of_the_Sample_Mean.html): The Variability of the Sample Mean
* [CIT 14.6](https://inferentialthinking.com/chapters/14/6/Choosing_a_Sample_Size.html): Choosing a Sample Size

You should complete this entire lab so that all tests pass and submit it to Gradescope by 11:59PM on the due date.

As usual, **run the cell** below to prepare the lab and the automatic tests.

```python
import babypandas as bpd
import numpy as np
np.set_printoptions(legacy='1.21')

import matplotlib.pyplot as plt
plt.style.use('seaborn-v0_8-colorblind')
plt.rcParams['figure.figsize'] = (10, 5)

import otter
grader = otter.Notebook('lab5.ipynb')
```

## 1. Comparing Grades Using Standard Units

Two of your friends, Kathleen and Eric, just took their midterms. Kathleen took her BILD 1 midterm and Eric took his MATH 18 midterm. Kathleen received a **B+** on her midterm (87%) and Eric received an **A-** (92%). Kathleen claims that while she received a lower grade on her midterm, she actually did better (relative to the rest of the class) than Eric. Eric disagrees. Knowing that you are taking DSC 10, your two friends come to you to settle their argument. 

They show you two DataFrames: `bild_midterm` and `math_midterm` that represent the grades for their classes. Both exams are out of 100 points.  Each DataFrame has a column called `'Student'` with student ID numbers and `'Score'` with the midterm scores.

_Note:_ You do ***not*** need to make any changes to the below cell. It is for you to visualize the two datasets.

```python
# Kathleen's exam
bild_midterm = bpd.read_csv("data/bild1_scores.csv")
bild_midterm.plot(y='Score', kind='hist', density=True, bins=range(0, 101, 1), ec='w', title='Distribution of BILD 1 Midterm Scores')
kathleen_score = 87
print("Kathleen's Score (BILD 1): " + str(kathleen_score))

# Eric's exam
math_midterm = bpd.read_csv("data/math18_scores.csv")
math_midterm.plot(y='Score', kind='hist', density=True, bins=range(0, 101, 1), ec='w', title='Distribution of MATH 18 Midterm Scores')
eric_score = 92
print("Eric's Score (MATH 18): " + str(eric_score))
```

You know that instead of comparing their actual scores, you should first convert their scores into **standard units**. Recall from [Lecture 16](https://dsc10.com/resources/lectures/lec16/lec16.html#Standard-units) that if $x$ is a numerical variable, and $x_i$ is one value of the variable, then

$$x_{i \: \text{(su)}} = \frac{x_i - \text{mean of $x$}}{\text{SD of $x$}}$$

represents $x_i$ in **standard units** – the number of standard deviations $x_i$ is above the mean.

To compute the midterm score in standard units for each friend, we need to:
1. Compute the **average** grade for each class. We will use the Series method `.mean()` to do this.
2. Compute the **standard deviation** (SD) of the midterm scores for each class. We *could* use `np.std`, but we will write our own function to do that.

Note that standard deviation is the square root of variance. So, we'll proceed by defining a function that computes the **variance** first. Recall that the variance is the mean squared deviation from the average:

$$\text{variance} = \frac{(\text{value}_1 - \text{average})^2 + (\text{value}_2 - \text{average})^2 +...+ (\text{value}_n - \text{average})^2}{n}$$
where $n$ is the number of values (e.g. number of exam scores, in our case).


**Question 1.1.** Complete the function `compute_variance`. It takes as input a *sequence of numbers* (`data`), which can be either a Series or an array, and returns the variance as a single number. 

Then, use the `compute_variance` function to compute the variance of the two classes' midterm scores, and assign them to the two specified variable names.

**Do not** use `np.std` in your solution. Instead, use the above formula for variance as guidance.

```python
def compute_variance(data):
    # Convert the input to be of type array. Don't change this line!
    data = np.array(data)
    n = ...
    average = ...
    sum_squared_diff = ...
    variance = sum_squared_diff / n 
    return variance


bild_midterm_var = ...
print("Variance of BILD 1 midterm: " + str(bild_midterm_var))

math_midterm_var = ...
print("Variance of MATH 18 midterm: " + str(math_midterm_var))
```

```python
grader.check("q1_1")
```

**Question 1.2.** Now that we have a function that computes the variance, we want to write a function that computes the standard deviation. Fill in the missing code to complete the function `compute_sd`. It takes as input a *sequence of numbers* (`data`), which can be either a Series or an array, and returns the *standard deviation* as a single number. 

Then, use the `compute_sd` function to compute the standard deviation of scores of the two midterms.

_Hint:_ Your implementation of `compute_sd` should only take ***one line***, that involves both the `return` keyword and the function `compute_variance`.



```python
def compute_sd(data):
    ...

bild_midterm_sd = ...
print("Standard Deviation of BILD 1 midterm: " + str(bild_midterm_sd))

math_midterm_sd = ...
print("Standard Deviation of MATH 18 midterm: " + str(math_midterm_sd))
```

```python
grader.check("q1_2")
```

**Question 1.3.** Now that you can compute the standard deviation, you are equipped to write a function that converts a given score to standard units. Fill in the missing code to complete the function `compute_su`. It takes in a *score* (`score`), the *average score* (`avg`), and the *standard deviation* (`sd`), and returns the score in standard units.

Then, use the `compute_su` function to transform the scores earned by each friend into standard units.

**Warning**: Be careful with the order of operations!

```python
def compute_su(score, avg, sd):
    standard_units = ...
    return standard_units

kathleen_su = ...
print("Kathleen's Score in Standard Units: " + str(kathleen_su))

eric_su = ...
print("Eric's Score in Standard Units: " + str(eric_su))
```

```python
grader.check("q1_3")
```

**Question 1.4.** Kathleen's score *is* higher than Eric's score when we convert to standard units, which can be seen as evidence that she did better on her exam relative to her classmates than Eric did relative to his. 

Another way to measure their relative performances is to calculate, for both Kathleen and Eric individually, the proportion of students they scored higher than (or the same as). Comparing Kathleen's proportion to Eric's proportion will give us another way of measuring who did better relative to their classmates. Calculate Kathleen's proportion and Eric's proportion below. (This will require looking at both `bild_midterm` and `math_midterm`.)

```python
kathleen_proportion = ...
print("Kathleen's Proportion: " + str(kathleen_proportion))

eric_proportion = ...
print("Eric's Proportion: " + str(eric_proportion))
```

```python
grader.check("q1_4")
```

## 2. Chebyshev's Inequality and Normal Curves

Lets look at the histograms of the scores of the two midterms again.

```python
bild_midterm.plot(y='Score', kind='hist', density=True, bins=range(0, 101, 1), ec='w', title='Distribution of BILD 1 Midterm Scores');
math_midterm.plot(y='Score', kind='hist', density=True, bins=range(0, 101, 1), ec='w', title='Distribution of MATH 18 Midterm Scores');
```

**Question 2.1.** Which of the two graphs roughly resembles a normal curve? Assign the variable `q2_1` to either 1, 2, 3, or 4.

1. Only the first graph (distribution of BILD 1 midterm scores) is roughly normal.
2. Only the lower graph (distribution of MATH 18 midterm scores) is roughly normal.
3. Both graphs are roughly normal.
4. Neither graph is roughly normal.

Remember all normal curves have the following characteristics:

* The mean (average) is always in the center of a normal curve.
* A normal curve has only one mode (peak).

```python
q2_1 = ...
q2_1
```

```python
grader.check("q2_1")
```

**Question 2.2.**
By looking at the distribution of MATH 18 midterm scores above, rank the following values in order **from smallest to largest**.

1. The mean score.
2. The median score.
3. The most common score (the mode).

Set variable `q2_2` to a list containing the numbers 1, 2, 3 **in the appropriate order**. Don't compute any of these values manually!

```python
q2_2 = ...
q2_2
```

```python
grader.check("q2_2")
```

### Recap: Chebyshev's inequality

Chebyshev's inequality states that *no matter what the shape of the distribution* is, the proportion of the values that fall in the range

$$\text{mean} \pm z \text { Standard Deviations}$$
is at least 
$$1 - \frac{1}{z^{2}}$$

**It's important to note that these are lower bounds, not approximations:** 75% of the data is guaranteed to lie within plus or minus of 2 standard deviations of the mean, but 100% of the data might also lie within plus or minus 2 standard deviations of the mean. 

### On the other hand...
**If we know that our data forms a normal curve**, the standard deviation is even more informative.

|Percent in Range | All Distributions (via Chebyshev's Inequality) | Normal Distributions|
|---|---|---|
|$\text{mean} \pm 1 \ \text{SD}$ | $\geq 0\%$ | $\approx 68\%$ |
|$\text{mean} \pm 2\text{SDs}$ | $\geq 75\%$ | $\approx 95\%$ |
|$\text{mean} \pm 3\text{SDs}$ | $\geq 88\%$ | $\approx 99.73\%$ |

**Note that for a normal distribution, the numbers in the last column of the DataFrame above are approximations, not lower bounds.**  
* If the distribution is perfectly normal, then 68% of the data (not more, not less) will lie between plus and minus one standard deviation of the mean. 
* Additionally because a normal curve is symmetric, we know that 34% of the data lies between the mean and the mean plus one standard deviation. 
 


**Question 2.3.**  Kathleen, who is Pre-Med, really wanted to score in the top 5% of the class. But before taking the exam, she did not know if the scores would be normally distributed or not. 

Without making any assumptions about the distribution of scores, how many standard deviations above the mean would she have needed to score to **guarantee** that she fell in the top 5% of the the class? Set the variable `q2_3` to either 1, 2, 3, or 4, depending on your answer. 


1. Kathleen would need to score roughly 4.5 standard deviations above the average to guarantee being in the top five percent. Using Chebyshev's inequality, setting $z = \sqrt{20} \approx 4.5$  gives that 95% of the data will lie between plus or minus 4.5 SDs. If Kathleen scores above 4.5 SDs, then she is guaranteed to have scored better than 95% of the other students. 

2. Kathleen would need to score above 2 SDs. Since 95% of the data falls between plus or minus 2 SDs, if Kathleen scores above 2 SDs, she is guaranteed to score above 95% of the class. 

3. Kathleen would need to score slightly less than 2 SDs. 50% of the class will have scored below the average. Which means that if Kathleen scores 2 standard deviations above the average she'll have scored higher than 50% + (95% / 2) = 97.5%. 

4. No matter how many standard deviations above the mean Kathleen scores, there is no guarantee that she will score in the top 5% of the class. 

```python
q2_3 = ...
q2_3
```

```python
grader.check("q2_3")
```

**Question 2.4.** Now, assuming that the scores for the exam will be normally distributed (as many exams are), what is the minimum number of standard deviations above the mean Kathleen would have needed to score to **guarantee** that she fell in the top 5% of the class? Set variable `q2_4` to either 1, 2, 3, or 4, depending on your answer. 


1. Kathleen would need to score roughly 4.5 standard deviations above the average to guarantee being in the top five percent. Using Chebyshev's inequality, setting $z = \sqrt{20} \approx 4.5$  gives that 95% of the data will lie between plus or minus 4.5 SDs. If Kathleen scores above 4.5 SDs, then she is guaranteed to have scored better than 95% of the other students. 

2. Kathleen would need to score above 2 SDs. Since 95% of the data falls between plus or minus 2 SDs, if Kathleen scores above 2 SDs, she is guaranteed to score above 95% of the class. 

3. Kathleen would need to score slightly less than 2 SDs. 50% of the class will have scored below the average. Which means that if Kathleen scores 2 standard deviations above the average she'll have scored higher than 50% + (95% / 2) = 97.5% of the class.

4. No matter how many standard deviations above the mean Kathleen scores, there is no guarantee that she will score in the top 5% of the class.

```python
q2_4 = ...
q2_4
```

```python
grader.check("q2_4")
```

Kathleen and Eric thank you for your analysis 👋, and go on their way to start studying for their finals, which are just around the corner. 💀

## 3. Choosing Sample Size

A new class is being offered at UCSD and the administration wants to know how many students will be taking the class so they know how big of a classroom it will need. To take the class, a student must have satisfied the prerequisites first. 

The administration knows there are 900 students who have completed the prerequisites and are eligible to take the class, but they don't know how many will actually enroll. They don't have the resources to ask each eligible student whether they are going to enroll, so they decide to ask a sample of the students, but they don't know how many students to ask. They want the width of their confidence interval to be **at most 10 students**. 

For example, if the results of their sample concluded that with 95% confidence between 200 and 210 students would take the class, the administration would be happy with that sample. However if the results of the sample concluded that with 95% confidence between 200 and 300 students would take the class, the sample would not have been informative enough because that range is too wide. We are going to help determine how big of a sample the administration should take. 

The population parameter we are interested in measuring is the proportion of eligible students who will take the class. We will estimate this using a sample statistic, the proportion of eligible students in the sample who plan to take the class. 

So **where do we start**?

The Central Limit Theorem tells us that regardless of the distribution of our population, the distribution of the sample mean or sum will always be roughly normal. Fortunately, our sample statistic (the proportion of eligible students who will take the class) is also a sample mean, because proportions are just means of 0s and 1s. Let's run a simulation to see this for ourselves. 

Below is the data for the whole population of eligible students. (If the administration had the resources to ask every student whether they were going to take the class, this is what they would see. `0` means they won't take the class and `1` means they will.)

```python
population = bpd.read_csv("data/population.csv")
population
```

```python
population.plot(y='Planning on taking', kind='hist', density=True, bins=np.arange(-0.5, 2.5, 1), ec='w');
plt.xticks([0, 1], [0, 1]);
```

**Question 3.1.** Below is partially implemented code to run a simulation. The simulation will repeatedly take samples of size `sample_size` (without replacement) from `population` and calculate the proportion of students who plan on taking the class. Fill in the missing parts. 



```python
def simulation(population, num_iterations, sample_size):
    results = np.array([])
    for i in np.arange(num_iterations):
        sampled = ...
        proportion_taking_class = ...
        results = ...
        
    bpd.DataFrame(data=results, columns=["Proportion"]).plot(kind='hist', 
                                                             y='Proportion', 
                                                             density=True,
                                                             ec='w',
                                                             bins=np.arange(0, 1, 1/(sample_size+1)),
                                                             title=f'Distribution of Sample Proportions (sample size = {sample_size})');
    plt.xlim(0, 1);
```

```python
grader.check("q3_1")
```

Run the cell below to see the empirical distribution of 10000 simulated sample proportions with a sample size of 40.

```python
simulation(population, 10000, 40)
```

**Question 3.2.** Does the distribution of the sample proportion look more like a normal curve or more like the population distribution? 

1. More like a normal curve.

2. More like the original population.

```python
q3_2 = ...
q3_2
```

```python
grader.check("q3_2")
```

We also know that as we increase the sample size, the standard deviation of our sample proportion's distribution will decrease. Again we decide to run a simulation to double check. Run the following cell to see how the distribution of the sample proportion changes as we increase the size of our sample. **It might take a while to run.** 

```python
simulation(population, 10000, 40)
simulation(population, 10000, 120)
simulation(population, 10000, 360)
```

This trend is expressed by the formula
$$ \text{SD of Distribution of Possible Sample Means} = \frac{\text{Population SD}}{\sqrt{\text{sample size}}}$$

Since proportions are means, we can use this formula to find the sample size we need to get a desired standard deviation of the sample proportion, and thus a certain confidence interval for that sample proportion. However, before taking our sample, we don't have any way of knowing the standard deviation of our population. Lecture 18, the [CIT](https://inferentialthinking.com/chapters/14/6/Choosing_a_Sample_Size.html) textbook, and Homework 5 include some ways to get around this problem; here we will use the actual population standard deviation. 

**Question 3.3.** The administration wants the confidence interval to have a width of **10 students**, but we have been calculating the proportion of eligible students who are planning on taking the class. Fortunately, there is a relationship between a count and a proportion that will allow us to convert between these.

Using the total number of students who are eligible to take the class, determine what proportion of that number equals 10 students, and save the result as `width_as_proportion`. For example, if 500 students are eligible to take the class, then 10 students as a proportion is 0.02.

```python
num_eligible_students = ...
print('Number of eligible students:', num_eligible_students)
width_as_proportion = ...
print('Desired confidence interval width, as a proportion:', width_as_proportion)
```

```python
grader.check("q3_3")
```

**Question 3.4.** Now let's calculate the standard deviation that our distribution of sample proportions would need to have in order for our 95% confidence interval to have a width of `width_as_proportion`. Set the variable `target_sd` to this value. 

Remember that for a normal distribution, 95% of the data lies between *plus and minus* 2 SDs of the mean. 

```python
target_sd = ...
target_sd
```

```python
grader.check("q3_4")
```

**Question 3.5.** We also need to calculate the standard deviation of the total population. Calculate this value using the`compute_sd` function that you wrote earlier and store it in the variable `population_sd`.

```python
population_sd = ...
population_sd
```

```python
grader.check("q3_5")
```

**Question 3.6.** Now calculate the required **sample size** and store your result as `req_sample_size`. Recall that

$$ \text{SD of Distribution of Possible Sample Means} = \frac{\text{Population SD}}{\sqrt{\text{sample size}}}$$

You have already calculated $\text{SD of Distribution of Possible Sample Means}$, in Question 3.4, and $\text{Population SD}$, in Question 3.5.

```python
req_sample_size = ...
req_sample_size
```

```python
grader.check("q3_6")
```

**Question 3.7.** Our required sample size is bigger than our entire population. For each part, say whether it is `True` or `False`.

1. The administration will have to settle for a wider interval to get 95% confidence.
2. Sampling with replacement will be a feasible way to determine the information the administration needs.
3. The administration will have to settle for a lower degree of confidence to get an interval of width 10.
4. We should increase the size of the population until the sample size is smaller than the size of the population.

Set each variable below to either `True` or `False`.


```python
statement_1 = ...
statement_2 = ...
statement_3 = ...
statement_4 = ...
```

```python
grader.check("q3_7")
```

# Finish Line 🏁

Congratulations! You are done with Lab 5.

<hr style="color:Maroon;background-color:Maroon;border:0 none; height: 3px;">

**Citations:** Did you use any generative artificial intelligence tools to assist you on this assignment? If so, please tell us which tools you used by setting `gen_ai` to a **list** of numbers. Include all tools used, from the options below:

1. AI Tutor, in Tutor Mode
2. AI Tutor, in ChatGPT Mode
3. Any other outside generative AI tool, such as Gemini, Claude, ChatGPT, etc.

If you didn't use any such tools, set `gen_ai` to an empty list, `[]`.

```python
gen_ai = ...
```

```python
grader.check("citations")
```

<hr style="color:Maroon;background-color:Maroon;border:0 none; height: 3px;">

To submit your assignment:

1. Select `Kernel -> Restart Kernel and Run All Cells` to ensure that you have executed all cells, including the test cells.
2. Read through the notebook to make sure everything is fine and all tests passed.
3. Run the cell below to run all tests, and make sure that they all pass.
4. Download your notebook as a `.ipynb` file using `File -> Download`, then upload your notebook  to Gradescope under "Lab 5".
5. **Stick around while the Gradescope autograder grades your work. Make sure you see that all tests have passed on Gradescope.**
6. Check that you have a confirmation email from Gradescope and save it as proof of your submission.

```python
# For your convenience, you can run this cell to run all the tests at once!
grader.check_all()
```

Tests for citations (points: None)
```
>>> bool(isinstance(gen_ai, list))
True
```
failure_message: In Python, a list has square brackets and comma-separated entries, such as [1, 2, 3]


Tests for q1_1 (points: None)
```
>>> np.isclose(bild_midterm_var, 173.20529599999983)
True
```
```
>>> np.isclose(math_midterm_var, 84.26559999999982)
True
```
```
>>> np.isclose(compute_variance(np.array([1, 2, 3, 4, 5])), 2.0)
True
```


Tests for q1_2 (points: None)
```
>>> np.isclose(bild_midterm_sd, 13.160748306992268)
True
```
```
>>> np.isclose(math_midterm_sd, 9.179629622157956)
True
```
```
>>> np.isclose(compute_sd(np.array([1, 2, 3, 4, 5])), 1.4142135623730951)
True
```


Tests for q1_3 (points: None)
```
>>> np.isclose(kathleen_su, 1.6676863266459616)
True
```
```
>>> np.isclose(eric_su, 0.38345773684630546)
True
```


Tests for q1_4 (points: None)
```
>>> kathleen_proportion == 0.954
True
```
```
>>> eric_proportion == 0.582
True
```


Tests for q2_1 (points: None)
```
>>> 0 < q2_1 < 5
True
```
```
>>> q2_1 != 2
True
```
```
>>> q2_1 != 3
True
```
```
>>> q2_1 != 4
True
```


Tests for q2_2 (points: None)
```
>>> q2_2 != []
True
```
```
>>> q2_2 != [2, 1, 3]
True
```
```
>>> q2_2 != [2, 3, 1]
True
```
```
>>> q2_2 != [1, 3, 2]
True
```
```
>>> q2_2 != [3, 1, 2]
True
```
```
>>> q2_2 != [3, 2, 1]
True
```


Tests for q2_3 (points: None)
```
>>> 0 < q2_3 < 5
True
```
```
>>> q2_3 != 2
True
```
```
>>> q2_3 != 3
True
```
```
>>> q2_3 != 4
True
```


Tests for q2_4 (points: None)
```
>>> 0 < q2_4 < 5
True
```
```
>>> q2_4 != 2
True
```
```
>>> q2_4 != 1
True
```
```
>>> q2_4 != 4
True
```


Tests for q3_1 (points: None)
```
>>> callable(simulation)
True
```


Tests for q3_2 (points: None)
```
>>> 0 < q3_2 < 3
True
```
```
>>> q3_2 != 2
True
```


Tests for q3_3 (points: None)
```
>>> np.isclose(width_as_proportion, 0.011111111111111112)
True
```


Tests for q3_4 (points: None)
```
>>> np.isclose(target_sd, 0.002777777777777778)
True
```


Tests for q3_5 (points: None)
```
>>> np.isclose(population_sd, 0.4582575694955866)
True
```


Tests for q3_6 (points: None)
```
>>> np.isclose(req_sample_size, 27216)
True
```


Tests for q3_7 (points: None)
```
>>> statement_1 == True
True
```
```
>>> statement_2 == False
True
```
```
>>> statement_3 == True
True
```
```
>>> statement_4 == False
True
```
